# Brève introduction à Python
----

Python est un language de programmation avec certaines spécificités vis à vis de languages comme C, Fortran, BASIC, PHP, etc...

- C'est un language **interpreté** donc pas de processus de compilation avant son execution comme en C ou Portran, Pascal. 
- Utilisation **interactive** ou sous forme de **scripts** (enchainement de commandes) exécutables.
- C'est un language **open-source**
- Il est **multi-plateforme** : Windows, Linux/Unix, MacOS X,
- Il est doté d'une **syntaxe claire, explicite et très lisible**,
- C'est un language possédant **un grand nombre de bibliothèques de grande qualité** pour des applications diverses, en allenat du web au calcul scientifique.
- Il s'interface très bien avec d'autres languages dont en particulier C and C++.
- C'est un language orienté-objet avec **typage dynamique** (la même variable peut contenir des objets de différents types durant l'execution du script)

## Premiers pas
---

```python
print("Bienvenue à Toulouse")
```

Output:
```
Bienvenue à Toulouse
```

```python
a = 3
b = a*2
print(b, type(b))
```

Output:
```
6 <class 'int'>
```

```python
b = "Python est super cool"
print(b, type(b))
```

Output:
```
Python est super cool <class 'str'>
```

```python
c = 'une chaine de caractères'
print(c+c)
print(c*3)
```

Output:
```
une chaine de caractèresune chaine de caractères
une chaine de caractèresune chaine de caractèresune chaine de caractères
```

```python

```

## Les types de bases
-----------
Les types de bases sont les entiers, flottants, complexes, booléens

```python
i = 1 # entier
f = 2.4 # flottant
c = 1+2j # complexe
print(c+i+f, type(c))
print(c.real, c.imag, c.conjugate())
```

```python
test = (i < 0)
print(test, type(test))
```

## Les opérations de base
-----------
- Les opérations disponibles sont : +, -, x, /, %(modulo), // (division entière), ** (puissance)

- Les conversions de types de données (le casting)

```python
a = 2+3**3
b = a/3
c = a//3
d = a%3
print(a, round(b,2), c, d)
```

```python
e = float(a)
print(e//3)
```

Levée d'une **exception**  par Python si un problème de conversion survient.

```python
f = str(a)
print(f//3) 
---------------------------------------------------------------------------
TypeError                                 Traceback (most recent call last)
<ipython-input-3-ce643e6b3450> in <module>()
      1 f = str(a)
----> 2 print(f//3)

TypeError: unsupported operand type(s) for //: 'str' and 'int'
---------------------------------------------------------------------------        
```  

<div class='alert alert-success' role='alert'>
    <h5 class='alert-heading'>Remarque :</h5>    
    <small>Python permet de capturer les exceptions et permettre un traitement spécifique si besoin.</small>
</div>

```python
f = str(a)
try:
    print(f//3)    
except TypeError as e:
    print(e)
```

## Les conteneurs
-----------


### Les listes
-----------
- Collection **ordonnée** et **modifiable** d'objets de type différent potentiellement
- Accès aux éléments par indexation. L'index du premier élément est **zéro**

```python
# une  liste vide
liste_1 = []
# une liste avec des objets variés (entier, caractère, liste)
liste_2 = [1, 'a', [2,3]]
print(liste_1, liste_2, liste_2[1])
```

#### Ajout d'éléments à une liste
---------

```python
# ajout d'un élément en fin de liste
liste_1.append(1)
# ajout d'une liste d'éléments en fin de liste
liste_1.extend([2,3])
# ajout d'une autre liste d'éléments
liste_1 = liste_1 + [4, 5]

liste_2 = liste_2*2

print("liste_1: {}\nliste_2: {}".format(liste_1,liste_2))
```

#### Les méthodes/fonctions associées 
------------
Les listes sont des objets possédant des méthodes d'ajout, de suppression d'élements, de tri, ...

Pour accèder à ces fonctions, dans une cellule, utilisez la touche <kbd>Tab</kbd> précédée du point comme cela:

```python
    liste_1.<tab>
```
et vous obtenez la liste des méthodes associées à l'objet liste_1.

```python
liste_1.pop(3)
```

```python
# On peut modifier les éléments d'une liste
liste_1[0] = 0
print(liste_1)
```

```python
try:
    liste_1[6]
except IndexError as e:
    print(e)
    print("Tu depasses les bornes !")
print("mais ce n'est pas grave")    
```

#### Parcourir une liste
------------

```python
# accès à un élément de la liste
print("premier élément {}, dernier élément {}".format(liste_1[0],
                                                      liste_1[-1]))
# on peut prendre une tranche
print(liste_1[0:3])
# on peut aussi inverser la liste facilement
print(liste_1[::-1])
# parcourir la liste avec un incrément fixé
print(liste_1[::2])
# et même à rebours
print(liste_1[::-2])
```

#### Tri d'une liste
---------

```python
liste_a_trier = [23, 3, 0, 10, 45]
liste_triee = sorted(liste_a_trier)
liste_triee_inverse = sorted(liste_a_trier, reverse=True)
print(liste_triee, liste_triee_inverse)
```

```python
try:
    liste_a_trier = [23, 3, 0, 10,'a', 45, 0+2j]
    print(sorted(liste_a_trier))
except TypeError as e:
    print(e)
```

### Les tuples
-----------
- Collection **ordonnée** et **non modifiable** d'objets de type différent potentiellement
- Accès aux éléments par indexation

```python
# un tuple 
tuple_1 = (1, 'a', [2,3])
print(tuple_1)
```

La modification du tuple n'est pas permise

```python
# ajout d'un élément
tuple_1[0]
try:
    tuple_1[0] = 2
except TypeError as e:
    print(e)
    print("Et oui, modification impossible!")
```

<p class='alert alert-info'>Les tuples ne sont pas modifiables (objets immuables).</p> Ils peuvent servir pour stocker des coordonnées dont on sait qu'on ne modifiera pas les valeurs.

### Les dictionnaires
-----------
Les dictionnaires sont des conteneurs liant des **clés ordonnées** à des **objets**.

```python
# création d'un dictionnaire vide
dico_1 = {}
# création d'un dictionnaire contenant 2 objets  
dico_1={'p1':(0,1),'p2':(2,4)}
# ajout d'une clé/valeur
dico_1['p3'] = (4,6)
dico_1['p4'] = (4,6)
# On peut vérifier que le dictionnaire ne conserve pas l'ordre de création des paires
print('le dico:', dico_1, '\np1:',dico_1['p1'], '\nles clés:',dico_1.keys())
```

```python
try:
    print(dico_1['p6'])
except KeyError as e:
    print(e)
    print('clé absente du dictionnaire')    
```

### Les chaines de caractères
-----------

Les chaines peuvent être créer de plusieures façons:

```python
# un guillement simple
chaine1 = 'chaine1'
# des guillemets doubles
chaine2 = "chaine2"
# des guillemets triples
chaine3 = '''chaine3
tient
sur deux lignes'''
print(chaine1, chaine2)
print(chaine3)
```

Si la chaine doit contenir des guillemets, il faut encadrer la chaine par des guillemets doubles.

```python
print("il s'agit de ne pas se tromper.")
```

#### L'accès et modification des éléments d'une chaine 
-----------

```python
print(chaine3+'\n')
print(chaine3[0:6])
try:    
    chaine3[6]=4
except TypeError as e:
    print(e)
    print(chaine3.replace('3','1',1))
```

#### Le formatage de chaine
-----------
Il existe plusieurs mécanisme pour formater les chaines de caractères. Par exemple:

```python
chaine_formatee = "Exemple de formatage de {} en python ".format(chaine1)
print(chaine_formatee)
ex2 = "la valeur réelle {0}, arrondie à 2 décimale {0:.2f} et en notation scientifique {0:.2e}".format(0.5465646)
print(ex2)
```

## Contrôles des flux
--------
<p class='alert alert-danger'>Les <b>blocs</b> de codes sont délimités par leur <b>indentation</b> (4 espaces).</p>

### Test if/elif/else
-------

```python
chaine_a_tester = 'abcdef'
if len(chaine_a_tester) > 5:
    print('la longueur de la chaine est supérieure à 5')
elif len(chaine_a_tester) <= 4:
    print('la longueur de la chaine est inférieure à 5')
else:
    print('la longueur de la chaine est égale à 5')
```

### Les boucles 
-------

#### for/range

```python
'''la range() renvoie un objet sur lequel
   on peut réaliser itérer des itérations'''
for i in range(4):
    print(i)
```

```python
# on peut aussi itérer sur une liste, un tuple
for el in ['super', 'génial', 'cool']:
    print("l'epu est est", el,'!')
```

#### While
-------

```python
z = 1 + 1j
while abs(z) < 100:
    z = z**2 + 1    
print(z)
```

#### Break et continue

```python
z = 1 + 1j
# on arrete la boucle 
while abs(z) < 100:
    if z.imag < 0:
        break
    z = z**2 + 1
print(z)
```

```python
a = [1, 0, 2, 4]
for element in a:
    if element == 0:
        continue
    print(1 / element)
```

## Fonctions
----------
Les fonctions sont définies par le mot clé ```def```. Notez l'indentation de 4 espaces dans le corps de la fonction.

```python
def ma_somme(parametre1, parametre2):
    variable_local = parametre1+parametre2
    return variable_local

print('retour de ma fonction:', ma_somme(1,2))
```

Notez que vous n'avez pas accès aux variables locales définies dans le corps de la fonction en dehors de cette fonction.

```python
try:
    print(variable_local)
except NameError as e:
    print(e)
```

## Pythonneries utiles
--------

### Compréhension de listes
-------------
Une des marques de fabrique de Python est d'être un language clair et concis notamment à l'aide de la fonctionnalité de *compréhension de liste* qui s'exprime :

```python
new_list = [function(item) for item in list if condition(item)]
```

Par exemple, imaginons que nous voulions filtrer une liste de valeurs pour ne garder que celles supérieures à 5 :

```python
ma_liste = [1, 3, 2, 5, 6, 2, 7, 6, 9]
new_liste = []
for element in ma_liste:
    if element > 5:
        new_liste.append(element)
print(new_liste)        
```

```python
new_liste = [element for element in ma_liste if element > 5]
print('Taaaaddaaa!', new_liste)
```

#### Exercice
----------
- A l'aide d'une *compréhension de liste*, mettez au carré les entiers $\in$ [1,20]:

<!-- grade_id: exo4, points: 0, grade: false, solution: true -->
```python
liste_depart = range(1,21)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: rep_exo4, points: 1, grade: true, solution: false -->
```python
for el,carre in zip(liste_depart,liste_carre):
    assert el**2 == carre 
```

### lambda()
-------------
Une autre particularité de python est la possibilité de définir des **fonctions anonymes**.
```python
fonction = lambda liste_arguments: expression
```

```python
somme = lambda x, y: x+y
print(somme(1,2))
```

#### Exercice
----------
- Transformez la fonction `f` suivante en fonction `lambda`:
```python
def f(x,y):
    if x > y:
        return x
    else:
        return y-x
```
- Appliquez la fonction `f` pour x=2 et y=3

<!-- grade_id: exo2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: rep_exo2, points: 1, grade: true, solution: false -->
```python
assert 'lambda' in f.__qualname__
assert f(2,3) == 1
```

### map()
-------------
Cette méthode execute une fonction sur chaque item d'un élément itérable.

```python
def fahrenheit(T):
    return ((float(9)/5)*T + 32)

temperatures_celsius = (36.5, 37, 37.5, 38, 39)
temperatures_farenheit = list(map(fahrenheit, temperatures_celsius))
print("les températures en °F:", temperatures_farenheit)
```

#### Exercice
----------

Utilisez une fonction anonyme (lambda) pour calculer les températures exprimées en °F en °C et utilisez cette fonction lambda dans la méthode map en prenant en paramètre de cette fonction la liste `temperatures_farenheit`.

<!-- grade_id: exo1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: rep_exo1, points: 1, grade: true, solution: false -->
```python
for i, temperature in enumerate(temperatures_celsius):
    assert mes_temp_celsius[i]-temperature < 1e-5
```